// Numbas version: exam_results_page_options
{"name": "Advanced Number Theory", "metadata": {"description": "", "licence": "Creative Commons Attribution-ShareAlike 4.0 International"}, "duration": 3600, "percentPass": 0, "showQuestionGroupNames": false, "shuffleQuestionGroups": false, "showstudentname": true, "question_groups": [{"name": "Group", "pickingStrategy": "all-ordered", "pickQuestions": 1, "questionNames": ["", "", "", "", "", ""], "variable_overrides": [[], [], [], [], [], []], "questions": [{"name": "B\u00e9zout", "extensions": [], "custom_part_types": [], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false}, "contributors": [{"name": "Julien Ugon", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/3575/"}], "tags": [], "metadata": {"description": "", "licence": "Creative Commons Attribution-ShareAlike 4.0 International"}, "statement": "", "advice": "", "rulesets": {}, "variables": {"a": {"name": "a", "group": "Ungrouped variables", "definition": "random(15..101)", "description": "", "templateType": "anything"}, "b": {"name": "b", "group": "Ungrouped variables", "definition": "random(15..101)", "description": "", "templateType": "anything"}, "coef": {"name": "coef", "group": "Ungrouped variables", "definition": "bezout(a,b)", "description": "", "templateType": "anything"}, "g": {"name": "g", "group": "Ungrouped variables", "definition": "gcd(a,b)", "description": "", "templateType": "anything"}}, "variablesTest": {"condition": "", "maxRuns": "100"}, "ungrouped_variables": ["a", "b", "coef", "g"], "variable_groups": [], "functions": {"bezout": {"parameters": [["a", "number"], ["b", "number"]], "type": "list", "language": "javascript", "definition": "function extendedEuclidean(b,a,s0,s1,t0,t1) {\n  if(a>b) return extendedEuclidean(a,b,t1,t0,s1,s0).reverse()\n  q = Math.floor(b/a)\n  r = b%a\n  if(r==0) return [s1,t1]\n  s2 = s0 - q*s1\n  t2 = t0 - q*t1\n  return extendedEuclidean(a,r,s1,s2,t1,t2)\n}\n\nreturn extendedEuclidean(a,b,1,0,0,1)"}}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "gapfill", "useCustomName": true, "customName": "B\u00e9zout Coefficients", "marks": 0, "showCorrectAnswer": true, "showFeedbackIcon": true, "scripts": {}, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "adaptiveMarkingPenalty": 0, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "prompt": "<p>Find the B&eacute;zout coefficients and the gcd of {a} and {b}:</p>\n<p>[[0]]$\\times \\var{a} +$ [[1]]$\\times \\var{b} =$ [[2]]</p>", "gaps": [{"type": "numberentry", "useCustomName": true, "customName": "a", "marks": 1, "showCorrectAnswer": true, "showFeedbackIcon": true, "scripts": {}, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "adaptiveMarkingPenalty": 0, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "minValue": "coef[0]", "maxValue": "coef[0]", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}, {"type": "numberentry", "useCustomName": true, "customName": "b", "marks": 1, "showCorrectAnswer": true, "showFeedbackIcon": true, "scripts": {}, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "adaptiveMarkingPenalty": 0, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "minValue": "coef[1]", "maxValue": "coef[1]", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}, {"type": "numberentry", "useCustomName": true, "customName": "gcd", "marks": 1, "showCorrectAnswer": true, "showFeedbackIcon": true, "scripts": {}, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "adaptiveMarkingPenalty": 0, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "minValue": "g", "maxValue": "g", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}], "sortAnswers": false}]}, {"name": "B\u00e9zout", "extensions": [], "custom_part_types": [], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false}, "contributors": [{"name": "Julien Ugon", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/3575/"}], "tags": [], "metadata": {"description": "", "licence": "Creative Commons Attribution-ShareAlike 4.0 International"}, "statement": "", "advice": "", "rulesets": {}, "variables": {"a": {"name": "a", "group": "Ungrouped variables", "definition": "random(15..101)", "description": "", "templateType": "anything"}, "b": {"name": "b", "group": "Ungrouped variables", "definition": "random(15..101)", "description": "", "templateType": "anything"}, "coef": {"name": "coef", "group": "Ungrouped variables", "definition": "bezout(a,b)", "description": "", "templateType": "anything"}, "g": {"name": "g", "group": "Ungrouped variables", "definition": "gcd(a,b)", "description": "", "templateType": "anything"}}, "variablesTest": {"condition": "", "maxRuns": "100"}, "ungrouped_variables": ["a", "b", "coef", "g"], "variable_groups": [], "functions": {"bezout": {"parameters": [["a", "number"], ["b", "number"]], "type": "list", "language": "javascript", "definition": "function extendedEuclidean(b,a,s0,s1,t0,t1) {\n  if(a>b) return extendedEuclidean(a,b,t1,t0,s1,s0).reverse()\n  q = Math.floor(b/a)\n  r = b%a\n  if(r==0) return [s1,t1]\n  s2 = s0 - q*s1\n  t2 = t0 - q*t1\n  return extendedEuclidean(a,r,s1,s2,t1,t2)\n}\n\nreturn extendedEuclidean(a,b,1,0,0,1)"}}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "gapfill", "useCustomName": true, "customName": "B\u00e9zout Coefficients", "marks": 0, "showCorrectAnswer": true, "showFeedbackIcon": true, "scripts": {}, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "adaptiveMarkingPenalty": 0, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "prompt": "<p>Find the B&eacute;zout coefficients and the gcd of {a} and {b}:</p>\n<p>[[0]]$\\times \\var{a} +$ [[1]]$\\times \\var{b} =$ [[2]]</p>", "gaps": [{"type": "numberentry", "useCustomName": true, "customName": "a", "marks": 1, "showCorrectAnswer": true, "showFeedbackIcon": true, "scripts": {}, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "adaptiveMarkingPenalty": 0, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "minValue": "coef[0]", "maxValue": "coef[0]", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}, {"type": "numberentry", "useCustomName": true, "customName": "b", "marks": 1, "showCorrectAnswer": true, "showFeedbackIcon": true, "scripts": {}, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "adaptiveMarkingPenalty": 0, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "minValue": "coef[1]", "maxValue": "coef[1]", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}, {"type": "numberentry", "useCustomName": true, "customName": "gcd", "marks": 1, "showCorrectAnswer": true, "showFeedbackIcon": true, "scripts": {}, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "adaptiveMarkingPenalty": 0, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "minValue": "g", "maxValue": "g", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}], "sortAnswers": false}]}, {"name": "Modular Inverse", "extensions": [], "custom_part_types": [], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false}, "contributors": [{"name": "Julien Ugon", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/3575/"}], "tags": [], "metadata": {"description": "", "licence": "Creative Commons Attribution-ShareAlike 4.0 International"}, "statement": "", "advice": "", "rulesets": {}, "variables": {"a1": {"name": "a1", "group": "Ungrouped variables", "definition": "random(15..101)", "description": "", "templateType": "anything"}, "b1": {"name": "b1", "group": "Ungrouped variables", "definition": "random(15..101)", "description": "", "templateType": "anything"}, "coef": {"name": "coef", "group": "Ungrouped variables", "definition": "bezout(a,b)", "description": "<p>Modular Inverse of a mod b</p>", "templateType": "anything"}, "g": {"name": "g", "group": "Ungrouped variables", "definition": "gcd(a1,b1)", "description": "", "templateType": "anything"}, "a": {"name": "a", "group": "Ungrouped variables", "definition": "a1/g", "description": "", "templateType": "anything"}, "b": {"name": "b", "group": "Ungrouped variables", "definition": "b1/g", "description": "", "templateType": "anything"}, "c": {"name": "c", "group": "Ungrouped variables", "definition": "mod(coef[0],b)", "description": "", "templateType": "anything"}, "niter": {"name": "niter", "group": "Ungrouped variables", "definition": "coef[2]", "description": "", "templateType": "anything"}}, "variablesTest": {"condition": "niter>3", "maxRuns": "100"}, "ungrouped_variables": ["a1", "b1", "coef", "g", "a", "b", "c", "niter"], "variable_groups": [], "functions": {"bezout": {"parameters": [["a", "number"], ["b", "number"]], "type": "list", "language": "javascript", "definition": "function extendedEuclidean(b,a,s0,s1,t0,t1,niter) {\n  if(a>b) return extendedEuclidean(a,b,s0,s1,t0,t1,niter)\n  q = Math.floor(b/a)\n  r = b%a\n  if(r==0) return [s1,t1,niter]\n  s2 = s0 - q*s1\n  t2 = t0 - q*t1\n  return extendedEuclidean(a,r,s1,s2,t1,t2,niter+1)\n}\n\n[x,y,n] = extendedEuclidean(a,b,1,0,0,1,0)\nif(a>b) return [x,y,n]\nelse return [y,x,n]"}}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "numberentry", "useCustomName": true, "customName": "Modular Inverse", "marks": 1, "showCorrectAnswer": true, "showFeedbackIcon": true, "scripts": {}, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "adaptiveMarkingPenalty": 0, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "prompt": "<p>What is the modular inverse of $\\var{a} \\bmod{\\var{b}}$? (Make sure the number is between 0 and {b-1}).</p>", "minValue": "c", "maxValue": "c", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}]}, {"name": "Modular Inverse", "extensions": [], "custom_part_types": [], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false}, "contributors": [{"name": "Julien Ugon", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/3575/"}], "tags": [], "metadata": {"description": "", "licence": "Creative Commons Attribution-ShareAlike 4.0 International"}, "statement": "", "advice": "", "rulesets": {}, "variables": {"a1": {"name": "a1", "group": "Ungrouped variables", "definition": "random(15..101)", "description": "", "templateType": "anything"}, "b1": {"name": "b1", "group": "Ungrouped variables", "definition": "random(15..101)", "description": "", "templateType": "anything"}, "coef": {"name": "coef", "group": "Ungrouped variables", "definition": "bezout(a,b)", "description": "<p>Modular Inverse of a mod b</p>", "templateType": "anything"}, "g": {"name": "g", "group": "Ungrouped variables", "definition": "gcd(a1,b1)", "description": "", "templateType": "anything"}, "a": {"name": "a", "group": "Ungrouped variables", "definition": "a1/g", "description": "", "templateType": "anything"}, "b": {"name": "b", "group": "Ungrouped variables", "definition": "b1/g", "description": "", "templateType": "anything"}, "c": {"name": "c", "group": "Ungrouped variables", "definition": "mod(coef[0],b)", "description": "", "templateType": "anything"}, "niter": {"name": "niter", "group": "Ungrouped variables", "definition": "coef[2]", "description": "", "templateType": "anything"}}, "variablesTest": {"condition": "niter>3", "maxRuns": "100"}, "ungrouped_variables": ["a1", "b1", "coef", "g", "a", "b", "c", "niter"], "variable_groups": [], "functions": {"bezout": {"parameters": [["a", "number"], ["b", "number"]], "type": "list", "language": "javascript", "definition": "function extendedEuclidean(b,a,s0,s1,t0,t1,niter) {\n  if(a>b) return extendedEuclidean(a,b,s0,s1,t0,t1,niter)\n  q = Math.floor(b/a)\n  r = b%a\n  if(r==0) return [s1,t1,niter]\n  s2 = s0 - q*s1\n  t2 = t0 - q*t1\n  return extendedEuclidean(a,r,s1,s2,t1,t2,niter+1)\n}\n\n[x,y,n] = extendedEuclidean(a,b,1,0,0,1,0)\nif(a>b) return [x,y,n]\nelse return [y,x,n]"}}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "numberentry", "useCustomName": true, "customName": "Modular Inverse", "marks": 1, "showCorrectAnswer": true, "showFeedbackIcon": true, "scripts": {}, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "adaptiveMarkingPenalty": 0, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "prompt": "<p>What is the modular inverse of $\\var{a} \\bmod{\\var{b}}$? (Make sure the number is between 0 and {b-1}).</p>", "minValue": "c", "maxValue": "c", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}]}, {"name": "Solving Equations", "extensions": [], "custom_part_types": [], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false}, "contributors": [{"name": "Julien Ugon", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/3575/"}], "tags": [], "metadata": {"description": "", "licence": "Creative Commons Attribution-ShareAlike 4.0 International"}, "statement": "", "advice": "", "rulesets": {}, "variables": {"a1": {"name": "a1", "group": "Ungrouped variables", "definition": "random(15..101)", "description": "", "templateType": "anything"}, "b1": {"name": "b1", "group": "Ungrouped variables", "definition": "random(15..101)", "description": "", "templateType": "anything"}, "coef": {"name": "coef", "group": "Ungrouped variables", "definition": "bezout(a,b)", "description": "<p>Result of the extended Euclidean algorithm: B&eacute;zout coefficients, and number of iterations.</p>", "templateType": "anything"}, "g": {"name": "g", "group": "Ungrouped variables", "definition": "gcd(a1,b1)", "description": "<p>greatest commond denominator of a and b.</p>", "templateType": "anything"}, "a": {"name": "a", "group": "Ungrouped variables", "definition": "a1/g", "description": "", "templateType": "anything"}, "b": {"name": "b", "group": "Ungrouped variables", "definition": "b1/g", "description": "", "templateType": "anything"}, "c": {"name": "c", "group": "Ungrouped variables", "definition": "mod(coef[0],b)", "description": "<p>Inverse of a mod(b)</p>", "templateType": "anything"}, "niter": {"name": "niter", "group": "Ungrouped variables", "definition": "coef[2]", "description": "", "templateType": "anything"}, "d": {"name": "d", "group": "Ungrouped variables", "definition": "random(2..(b-1))", "description": "", "templateType": "anything"}, "x": {"name": "x", "group": "Ungrouped variables", "definition": "mod(c*d,b)", "description": "", "templateType": "anything"}}, "variablesTest": {"condition": "niter>3", "maxRuns": "100"}, "ungrouped_variables": ["a1", "b1", "coef", "g", "a", "b", "c", "niter", "d", "x"], "variable_groups": [{"name": "Unnamed group", "variables": []}], "functions": {"bezout": {"parameters": [["a", "number"], ["b", "number"]], "type": "list", "language": "javascript", "definition": "function extendedEuclidean(b,a,s0,s1,t0,t1,niter) {\n  if(a>b) return extendedEuclidean(a,b,s0,s1,t0,t1,niter)\n  q = Math.floor(b/a)\n  r = b%a\n  if(r==0) return [s1,t1,niter]\n  s2 = s0 - q*s1\n  t2 = t0 - q*t1\n  return extendedEuclidean(a,r,s1,s2,t1,t2,niter+1)\n}\n\n[x,y,n] = extendedEuclidean(a,b,1,0,0,1,0)\nif(a>b) return [x,y,n]\nelse return [y,x,n]"}}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "numberentry", "useCustomName": true, "customName": "Modular Inverse", "marks": 1, "showCorrectAnswer": true, "showFeedbackIcon": true, "scripts": {}, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "adaptiveMarkingPenalty": 0, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "prompt": "<p>Solve for $x$ in $\\var{a}x\\equiv \\var{d} \\bmod{\\var{b}}$? so that $x$ is between 0 and {b-1}.</p>", "minValue": "x", "maxValue": "x", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}]}, {"name": "Solving Equations", "extensions": [], "custom_part_types": [], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false}, "contributors": [{"name": "Julien Ugon", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/3575/"}], "tags": [], "metadata": {"description": "", "licence": "Creative Commons Attribution-ShareAlike 4.0 International"}, "statement": "", "advice": "", "rulesets": {}, "variables": {"a1": {"name": "a1", "group": "Ungrouped variables", "definition": "random(15..101)", "description": "", "templateType": "anything"}, "b1": {"name": "b1", "group": "Ungrouped variables", "definition": "random(15..101)", "description": "", "templateType": "anything"}, "coef": {"name": "coef", "group": "Ungrouped variables", "definition": "bezout(a,b)", "description": "<p>Result of the extended Euclidean algorithm: B&eacute;zout coefficients, and number of iterations.</p>", "templateType": "anything"}, "g": {"name": "g", "group": "Ungrouped variables", "definition": "gcd(a1,b1)", "description": "<p>greatest commond denominator of a and b.</p>", "templateType": "anything"}, "a": {"name": "a", "group": "Ungrouped variables", "definition": "a1/g", "description": "", "templateType": "anything"}, "b": {"name": "b", "group": "Ungrouped variables", "definition": "b1/g", "description": "", "templateType": "anything"}, "c": {"name": "c", "group": "Ungrouped variables", "definition": "mod(coef[0],b)", "description": "<p>Inverse of a mod(b)</p>", "templateType": "anything"}, "niter": {"name": "niter", "group": "Ungrouped variables", "definition": "coef[2]", "description": "", "templateType": "anything"}, "d": {"name": "d", "group": "Ungrouped variables", "definition": "random(2..(b-1))", "description": "", "templateType": "anything"}, "x": {"name": "x", "group": "Ungrouped variables", "definition": "mod(c*d,b)", "description": "", "templateType": "anything"}}, "variablesTest": {"condition": "niter>3", "maxRuns": "100"}, "ungrouped_variables": ["a1", "b1", "coef", "g", "a", "b", "c", "niter", "d", "x"], "variable_groups": [{"name": "Unnamed group", "variables": []}], "functions": {"bezout": {"parameters": [["a", "number"], ["b", "number"]], "type": "list", "language": "javascript", "definition": "function extendedEuclidean(b,a,s0,s1,t0,t1,niter) {\n  if(a>b) return extendedEuclidean(a,b,s0,s1,t0,t1,niter)\n  q = Math.floor(b/a)\n  r = b%a\n  if(r==0) return [s1,t1,niter]\n  s2 = s0 - q*s1\n  t2 = t0 - q*t1\n  return extendedEuclidean(a,r,s1,s2,t1,t2,niter+1)\n}\n\n[x,y,n] = extendedEuclidean(a,b,1,0,0,1,0)\nif(a>b) return [x,y,n]\nelse return [y,x,n]"}}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "numberentry", "useCustomName": true, "customName": "Modular Inverse", "marks": 1, "showCorrectAnswer": true, "showFeedbackIcon": true, "scripts": {}, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "adaptiveMarkingPenalty": 0, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "prompt": "<p>Solve for $x$ in $\\var{a}x\\equiv \\var{d} \\bmod{\\var{b}}$? so that $x$ is between 0 and {b-1}.</p>", "minValue": "x", "maxValue": "x", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}]}]}], "allowPrinting": true, "navigation": {"allowregen": false, "reverse": true, "browse": true, "allowsteps": false, "showfrontpage": true, "showresultspage": "oncompletion", "navigatemode": "sequence", "onleave": {"action": "none", "message": ""}, "preventleave": true, "startpassword": ""}, "timing": {"allowPause": false, "timeout": {"action": "none", "message": ""}, "timedwarning": {"action": "none", "message": ""}}, "feedback": {"showactualmark": false, "showtotalmark": false, "showanswerstate": false, "allowrevealanswer": false, "advicethreshold": 0, "intro": "", "end_message": "", "reviewshowscore": true, "reviewshowfeedback": true, "reviewshowexpectedanswer": true, "reviewshowadvice": true, "feedbackmessages": []}, "diagnostic": {"knowledge_graph": {"topics": [], "learning_objectives": []}, "script": "diagnosys", "customScript": ""}, "contributors": [{"name": "Julien Ugon", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/3575/"}], "extensions": [], "custom_part_types": [], "resources": []}